# Assignment 2

For this assignment you'll be looking at 2017 data on immunizations from the CDC. Your datafile for this assignment is in [assets/NISPUF17.csv](assets/NISPUF17.csv). A data users guide for this, which you'll need to map the variables in the data to the questions being asked, is available at [assets/NIS-PUF17-DUG.pdf](assets/NIS-PUF17-DUG.pdf). **Note: you may have to go to your Jupyter tree (click on the Coursera image) and navigate to the assignment 2 assets folder to see this PDF file).**

## Question 1
Write a function called `proportion_of_education` which returns the proportion of children in the dataset who had a mother with the education levels equal to less than high school (<12), high school (12), more than high school but not a college graduate (>12) and college degree.

*This function should return a dictionary in the form of (use the correct numbers, do not round numbers):* 
```
    {"less than high school":0.2,
    "high school":0.4,
    "more than high school but not college":0.2,
    "college":0.2}
```

```python
import pandas as pd
df=pd.read_csv('assets/NISPUF17.csv')
#keep_columns=['SEQNUMC','EDUC1']
#df=df[keep_columns] #only keep
df=df.rename(columns={'SEQNUMC':'Children_ID','EDUC1':'Edu'})
#df.head()
```

```python
df3={"less than high school":0,
    "high school":0,
    "more than high school but not college":0,
    "college":0}
Mother_Edu=df['Edu']
import numpy as np
n=len(Mother_Edu)
df3['less than high school']=np.sum(Mother_Edu==1)/n
df3['high school']=np.sum(Mother_Edu==2)/n
df3['more than high school but not college']=np.sum(Mother_Edu==3)/n
df3['college']=np.sum(Mother_Edu==4)/n

df3
```

Output:
```
{'less than high school': 0.10202002459160373,
 'high school': 0.172352011241876,
 'more than high school but not college': 0.24588090637625154,
 'college': 0.47974705779026877}
```

```python
#df2=df.groupby(["Mother_Edu"])["Children_ID"].count()
#df2
```

Output:
```
Mother_Edu
1     2904
2     4906
3     6999
4    13656
Name: Children_ID, dtype: int64
```

<!-- grade_id: cell-eea16d020eb52ae7, points: 0, grade: false, solution: true -->
```python
def proportion_of_education():
    # your code goes here
    # YOUR CODE HERE
    import pandas as pd
    import numpy as np
    df=pd.read_csv('assets/NISPUF17.csv')
    df=df.rename(columns={'SEQNUMC':'Children_ID','EDUC1':'Edu'})
    df3={"less than high school":0,
    "high school":0,
    "more than high school but not college":0,
    "college":0}
    Mother_Edu=df['Edu']

    n=len(Mother_Edu)
    df3['less than high school']=np.sum(Mother_Edu==1)/n
    df3['high school']=np.sum(Mother_Edu==2)/n
    df3['more than high school but not college']=np.sum(Mother_Edu==3)/n
    df3['college']=np.sum(Mother_Edu==4)/n
    return df3
    
    #raise NotImplementedError()
```

<!-- grade_id: cell-c0eeef201366f51c, points: 1, grade: true, solution: false -->
```python
assert type(proportion_of_education())==type({}), "You must return a dictionary."
assert len(proportion_of_education()) == 4, "You have not returned a dictionary with four items in it."
assert "less than high school" in proportion_of_education().keys(), "You have not returned a dictionary with the correct keys."
assert "high school" in proportion_of_education().keys(), "You have not returned a dictionary with the correct keys."
assert "more than high school but not college" in proportion_of_education().keys(), "You have not returned a dictionary with the correct keys."
assert "college" in proportion_of_education().keys(), "You have not returned a dictionary with the correct keys."

```

## Question 2

Let's explore the relationship between being fed breastmilk as a child and getting a seasonal influenza vaccine from a healthcare provider. Return a tuple of the average number of influenza vaccines for those children we know received breastmilk as a child and those who know did not.

*This function should return a tuple in the form (use the correct numbers:*
```
(2.5, 0.1)
```

```python
import pandas as pd
import numpy as np
df=pd.read_csv('assets/NISPUF17.csv')
```

```python
#checking index
cols=list(df.columns)
cols
```

Output:
```
['Unnamed: 0',
 'SEQNUMC',
 'SEQNUMHH',
 'PDAT',
 'PROVWT_D',
 'RDDWT_D',
 'STRATUM',
 'YEAR',
 'AGECPOXR',
 'HAD_CPOX',
 'AGEGRP',
 'BF_ENDR06',
 'BF_EXCLR06',
 'BF_FORMR08',
 'BFENDFL06',
 'BFFORMFL06',
 'C1R',
 'C5R',
 'CBF_01',
 'CEN_REG',
 'CHILDNM',
 'CWIC_01',
 'CWIC_02',
 'EDUC1',
 'FRSTBRN',
 'I_HISP_K',
 'INCPORAR',
 'INCPORAR_I',
 'INCPOV1',
 'INCQ298A',
 'LANGUAGE',
 'M_AGEGRP2',
 'MARITAL2',
 'MOBIL_I',
 'NUM_PHONE',
 'NUM_CELLS_HH',
 'NUM_CELLS_PARENTS',
 'RACE_K',
 'RACEETHK',
 'RENT_OWN',
 'SEX',
 'ESTIAP17',
 'EST_GRANT',
 'STATE',
 'D6R',
 'D7',
 'N_PRVR',
 'PROV_FAC',
 'REGISTRY',
 'VFC_ORDER',
 'HEP_BRTH',
 'HEP_FLAG',
 'P_NUHEPX',
 'P_NUHIBX',
 'P_NUHPHB',
 'P_NUMDAH',
 'P_NUMDHI',
 'P_NUMDIH',
 'P_NUMDTA',
 'P_NUMDTP',
 'P_NUMFLU',
 'P_NUMFLUL',
 'P_NUMFLUM',
 'P_NUMFLUN',
 'P_NUMHEA',
 'P_NUMHEN',
 'P_NUMHEP',
 'P_NUMHG',
 'P_NUMHHY',
 'P_NUMHIB',
 'P_NUMHIN',
 'P_NUMHION',
 'P_NUMHM',
 'P_NUMHS',
 'P_NUMIPV',
 'P_NUMMCN',
 'P_NUMMMR',
 'P_NUMMMRX',
 'P_NUMMMX',

```

```python
df['CBF_01']
```

Output:
```
0        1
1        2
2        2
3        2
4        1
        ..
28460    2
28461    1
28462    1
28463    2
28464    1
Name: CBF_01, Length: 28465, dtype: int64
```

```python
df['P_NUMFLU']
```

Output:
```
0        NaN
1        NaN
2        NaN
3        3.0
4        0.0
        ... 
28460    NaN
28461    NaN
28462    NaN
28463    NaN
28464    NaN
Name: P_NUMFLU, Length: 28465, dtype: float64
```

```python
#Keep CBF_01 and P_NUMFLUE columns =>df2
df2=df[['CBF_01','P_NUMFLU']]
df2
```

Output:
```
CBF_01  P_NUMFLU
0           1       NaN
1           2       NaN
2           2       NaN
3           2       3.0
4           1       0.0
...       ...       ...
28460       2       NaN
28461       1       NaN
28462       1       NaN
28463       2       NaN
28464       1       NaN

[28465 rows x 2 columns]
```

```python
#Rename
df2=df2.rename(columns={'CBF_01':'ReceiveBM','P_NUMFLU':'NumDose'})

#delete NaN value
#df2=df2.dropna(subset = ['NumDose'])
#df2
```

```python
#define table of Received BreastMilk
Yes_BM=df2[df2['ReceiveBM']==1]
#define table of None Received BreastMilk
No_BM=df2[df2['ReceiveBM']==2]
```

```python
Yes_BM[np.isnan(Yes_BM)] = 0
n_yes=len(Yes_BM)
AVG_FDose_BM=Yes_BM['NumDose'].sum()/n_yes
AVG_FDose_BM
```

Output:
```
1.0298833518816206
```

```python
No_BM[np.isnan(No_BM)] = 0
n_no=len(No_BM)
AVG_FDose_notBM=No_BM['NumDose'].sum()/n_no
AVG_FDose_notBM
```

Output:
```
0.7804161566707466
```

```python
a=(AVG_FDose_BM,AVG_FDose_notBM)
```

<!-- grade_id: cell-77f18c512324eabb, points: 0, grade: false, solution: true -->
```python
def average_influenza_doses():
    # YOUR CODE HERE
    # raise NotImplementedError()
    import pandas as pd
    import numpy as np
    df = pd.read_csv("assets/NISPUF17.csv", index_col=0)
    
    cbf_flu=df.loc[:,['CBF_01','P_NUMFLU']]
    
    
    cbf_flu1=cbf_flu[cbf_flu['CBF_01'] ==1].dropna()
    cbf_flu2=cbf_flu[cbf_flu['CBF_01'] ==2].dropna()
    
    flu1=cbf_flu1['P_NUMFLU'].values.copy()
    flu1[np.isnan(flu1)] = 0
    f1=np.sum(flu1)/len(flu1)
    
    flu2=cbf_flu2['P_NUMFLU'].values.copy()
    flu2[np.isnan(flu2)] = 0
    f2=np.sum(flu2)/len(flu2)
    
    aid =(f1,f2)
    return aid
average_influenza_doses()
```

```python
""" Wrong answer
def average_influenza_doses():
    # YOUR CODE HERE
    import pandas as pd
    import numpy as np
    df=pd.read_csv('assets/NISPUF17.csv')
    
    #Keep CBF_01 and P_NUMFLUE columns => df2
    df2=df[['CBF_01','P_NUMFLU']]
    
    #Rename
    df2=df2.rename(columns={'CBF_01':'ReceiveBM','P_NUMFLU':'NumDose'})
    
    #define table of Received BreastMilk
    Yes_BM=df2[df2['ReceiveBM']==1]
    
    #define table of None Received BreastMilk
    No_BM=df2[df2['ReceiveBM']==2]
    
    Yes_BM[np.isnan(Yes_BM)] = 0
    n_yes=len(Yes_BM)
    AVG_FDose_BM=Yes_BM['NumDose'].sum()/n_yes
    AVG_FDose_BM
    
    No_BM[np.isnan(No_BM)] = 0
    n_no=len(No_BM)
    AVG_FDose_notBM=No_BM['NumDose'].sum()/n_no
    AVG_FDose_notBM
    
    a=(AVG_FDose_BM,AVG_FDose_notBM)
    return a
   # raise NotImplementedError()
```

```python
average_influenza_doses()
```

Output:
```
(1.0298833518816206, 0.7804161566707466)
```

<!-- grade_id: cell-54a3ba6cff31caa7, points: 1, grade: true, solution: false -->
```python
assert len(average_influenza_doses())==2, "Return two values in a tuple, the first for yes and the second for no."

```

Output:
```
Input In [65], in average_influenza_doses()
     11 df2=df2.rename(columns={'CBF_01':'ReceiveBM','P_NUMFLU':'NumDose'})
     13 #define table of Received BreastMilk
---> 14 Yes_BM=df2[df2['ReceiveBM']==1].value()
     16 #define table of None Received BreastMilk
     17 No_BM=df2[df2['ReceiveBM']==2].value()

File /opt/conda/lib/python3.8/site-packages/pandas/core/generic.py:5902, in NDFrame.__getattr__(self, name)
   5895 if (
   5896     name not in self._internal_names_set
   5897     and name not in self._metadata
   5898     and name not in self._accessors
   5899     and self._info_axis._can_hold_identifiers_and_holds_name(name)
   5900 ):
   5901     return self[name]
-> 5902 return object.__getattribute__(self, name)

AttributeError: 'DataFrame' object has no attribute 'value'
```

## Question 3
It would be interesting to see if there is any evidence of a link between vaccine effectiveness and sex of the child. Calculate the ratio of the number of children who contracted chickenpox but were vaccinated against it (at least one varicella dose) versus those who were vaccinated but did not contract chicken pox. Return results by sex. 

*This function should return a dictionary in the form of (use the correct numbers):* 
```
    {"male":0.2,
    "female":0.4}
```

Note: To aid in verification, the `chickenpox_by_sex()['female']` value the autograder is looking for starts with the digits `0.0077`.

```python
import pandas as pd
import numpy as np
df=pd.read_csv('assets/NISPUF17.csv')

#Keep HAD_CPOX and P_NUMVRC columns => df2
df2=df[['SEX','HAD_CPOX','P_NUMVRC']]
df2
```

Output:
```
SEX  HAD_CPOX  P_NUMVRC
0        1         2       NaN
1        1         2       NaN
2        2         2       NaN
3        2         2       1.0
4        2         2       0.0
...    ...       ...       ...
28460    2         2       NaN
28461    2         2       NaN
28462    2         2       NaN
28463    2         2       NaN
28464    2         2       NaN

[28465 rows x 3 columns]
```

```python
#Check had_cpox column
print(df2['HAD_CPOX'].head(20))
```

Output:
```
0     2
1     2
2     2
3     2
4     2
5     2
6     2
7     2
8     1
9     2
10    2
11    2
12    2
13    2
14    2
15    2
16    2
17    2
18    2
19    2
Name: HAD_CPOX, dtype: int64
```

```python
#df with vacinced
df3=df2[df2['P_NUMVRC']>=1]

#df of male
male_df=df3[df3['SEX']==1]

#df of female
female_df=df3[df3['SEX']==2]
```

```python
#Male df calculation

CP_male=male_df[male_df['HAD_CPOX']==1]
n_CP_male=len(CP_male)

nonCP_male=male_df[male_df['HAD_CPOX']==2]
n_nonCP_male=len(nonCP_male)

ratio_male=n_CP_male/n_nonCP_male

ratio_male
```

Output:
```
0.009675583380762664
```

```python
#Female df calculation

CP_female=female_df[female_df['HAD_CPOX']==1]
n_CP_female=len(CP_female)

nonCP_female=female_df[female_df['HAD_CPOX']==2]
n_nonCP_female=len(nonCP_female)

ratio_female=n_CP_female/n_nonCP_female

ratio_female
```

Output:
```
0.0077918259335489565
```

<!-- grade_id: cell-a0a9e6fe67698006, points: 0, grade: false, solution: true -->
```python
def chickenpox_by_sex():
    # YOUR CODE HERE
    import pandas as pd
    import numpy as np
    df=pd.read_csv('assets/NISPUF17.csv')

    df2=df[['SEX','HAD_CPOX','P_NUMVRC']]
    
    #df with vacinced
    df3=df2[df2['P_NUMVRC']>=1]

    #df of male
    male_df=df3[df3['SEX']==1]

    #df of female
    female_df=df3[df3['SEX']==2]
    
    
    b={"male":0,"female":0}
    
    #Male df calculation

    CP_male=male_df[male_df['HAD_CPOX']==1]
    n_CP_male=len(CP_male)

    nonCP_male=male_df[male_df['HAD_CPOX']==2]
    n_nonCP_male=len(nonCP_male)

    b['male']=n_CP_male/n_nonCP_male
    
    #Female df calculation

    CP_female=female_df[female_df['HAD_CPOX']==1]
    n_CP_female=len(CP_female)

    nonCP_female=female_df[female_df['HAD_CPOX']==2]
    n_nonCP_female=len(nonCP_female)

    b['female']=n_CP_female/n_nonCP_female

    return b
    #raise NotImplementedError()
```

```python
chickenpox_by_sex()
```

Output:
```
{'female': 0.0077918259335489565, 'male': 0.009675583380762664}
```

<!-- grade_id: cell-c4f1714db100c865, points: 1, grade: true, solution: false -->
```python
assert len(chickenpox_by_sex())==2, "Return a dictionary with two items, the first for males and the second for females."

```

## Question 4
A correlation is a statistical relationship between two variables. If we wanted to know if vaccines work, we might look at the correlation between the use of the vaccine and whether it results in prevention of the infection or disease [1]. In this question, you are to see if there is a correlation between having had the chicken pox and the number of chickenpox vaccine doses given (varicella).

Some notes on interpreting the answer. The `had_chickenpox_column` is either `1` (for yes) or `2` (for no), and the `num_chickenpox_vaccine_column` is the number of doses a child has been given of the varicella vaccine. A positive correlation (e.g., `corr > 0`) means that an increase in `had_chickenpox_column` (which means more no’s) would also increase the values of `num_chickenpox_vaccine_column` (which means more doses of vaccine). If there is a negative correlation (e.g., `corr < 0`), it indicates that having had chickenpox is related to an increase in the number of vaccine doses.

Also, `pval` is the probability that we observe a correlation between `had_chickenpox_column` and `num_chickenpox_vaccine_column` which is greater than or equal to a particular value occurred by chance. A small `pval` means that the observed correlation is highly unlikely to occur by chance. In this case, `pval` should be very small (will end in `e-18` indicating a very small number).

[1] This isn’t really the full picture, since we are not looking at when the dose was given. It’s possible that children had chickenpox and then their parents went to get them the vaccine. Does this dataset have the data we would need to investigate the timing of the dose?

<!-- grade_id: cell-8afff07f564cf79a, points: 0, grade: false, solution: true -->
```python
def corr_chickenpox():
    import scipy.stats as stats
    import numpy as np
    import pandas as pd

    # this is just an example dataframe
    # df=pd.DataFrame({"had_chickenpox_column":np.random.randint(1,3,size=(100)),
    #                "num_chickenpox_vaccine_column":np.random.randint(0,6,size=(100))})
    
    df = pd.read_csv("assets/NISPUF17.csv", index_col=0)

    df=df[df['HAD_CPOX'].lt(3)].loc[:,['HAD_CPOX','P_NUMVRC']].dropna()
    df.columns=['had_chickenpox_column','num_chickenpox_vaccine_column']
    # here is some stub code to actually run the correlation
    corr, pval=stats.pearsonr(df["had_chickenpox_column"],df["num_chickenpox_vaccine_column"])
    return corr
```

```python
import scipy.stats as stats
import numpy as np
import pandas as pd

    # this is just an example dataframe
    # df=pd.DataFrame({"had_chickenpox_column":np.random.randint(1,3,size=(100)),
    #                "num_chickenpox_vaccine_column":np.random.randint(0,6,size=(100))})
df = pd.read_csv("assets/NISPUF17.csv", index_col=0)

df=df[df['HAD_CPOX'].lt(3)].loc[:,['HAD_CPOX','P_NUMVRC']].dropna()
df.columns=['had_chickenpox_column','num_chickenpox_vaccine_column']
    # here is some stub code to actually run the correlation
corr, pval=stats.pearsonr(df["had_chickenpox_column"],df["num_chickenpox_vaccine_column"])
corr
```

Output:
```
0.07044873460147986
```

<!-- grade_id: cell-73408733533a29a5, points: 1, grade: true, solution: false -->
```python
assert -1<=corr_chickenpox()<=1, "You must return a float number between -1.0 and 1.0."

```